# Численное Интегрирование (обязательно сделать хотя бы 1 из звёздочек)

# Квадратурные формулы Ньютона-Котеса.

Пусть некоторая функция $f(x)$ задана в узлах интерполяции $x_{i}=x_{0}+i h(i=1,2, \ldots, n)$ на отрезке $[a, b]$ таблицей своих значений:

\begin{array}{c|c|c|c|c}
x_{0}=a & x_{1} & x_{2} & \ldots & x_{n}=b \\
\hline y_{0}=f\left(x_{0}\right) & y_{1}=f\left(x_{1}\right) & y_{2}=f\left(x_{2}\right) & \ldots & y_{n}=f\left(x_{n}\right)
\end{array} 

По заданным значениям подынтегральной терполяционный многочлен Лагранжа. Для равноотстояших узлов интерноляционный многочлен Лагранжа. Для равноотстоящих узлов интерполяционный многочлен Лагранжа имеет вид  

$
L_{n}(x)=\sum_{i=0}^{n} \frac{(-1)^{n-1}}{i !(n-i) !} \frac{q(q-1) \ldots(q-n)}{q-1} y_{i}  
$

$
\text { где } q=\frac{x-x_{0}}{h}-\text { шаг интерполяции. } 
$

Заменим подынтегральную функцию $f(x)$ интерполяционным многочленом Лагранжа, тогда
$$
\int_{a}^{b} f(x) d x \approx \int_{x_{i}}^{x_{n}}\left[\sum_{i=0}^{n} \frac{(-1) n-1}{i !(n-i) !} \frac{q(q-1) \ldots(q-n)}{q-i} y_{i}\right] d x .
$$
Поменяем местами знак суммирования и интеграл и вынесем за знак интеграла постоянные коэффициенты:
$$
\int_{a}^{b} f(x) d x \approx \sum_{i=0}^{n} y_{i} \frac{(-1)^{n-1}}{i !(n-i) !} \int_{x_{0}}^{x_{n}} \frac{q(q-1) \ldots(q-n)}{q-i} d x .
$$
Поскольку $d q=\frac{d x}{h}$ и $h=\frac{b-a}{n}$, то, заменив пределы интегрирования, получим
$$
\begin{gathered}
\int_{a}^{b} f(x) d x \approx(b-a) \times \\
\times \sum_{i=0}^{n} y_{i} \frac{(-1)^{n-i}}{i !(n-i) !} \frac{1}{n} \int_{0}^{n} q(q-1) \ldots(q-i+1)(q-i-1) \ldots(q-n) d q .
\end{gathered}
$$

Положим

$$
H_{i}=\frac{(-1)^{n-i}}{i !(n-i) !} \frac{1}{n} \int_{0}^{n} q(q-1) \ldots(q-i+1)(q-i-1) \ldots(q-n) d q
$$

где $i=0,1,2, \ldots, n$.  

Числа $H_{i}$ называются коэффициентами Ньютона - Котеса. Они не зависят от вида функции $f(x)$ и являются функцией только $n$ (количества узлов интерполирования). Коэффициенты Ньютона - Котеса можно вычислить заранее для различного числа узлов интерполирования и свести в табл. 4.1.
Окончательный вид формулы Ньютон - Котеса:
$$
\int_{a}^{b} f(x) d x \approx(b-a) \sum_{i=0}^{n} y_{i} H_{i}
$$  

\begin{array}{|c|c|}
\hline n=1 & H_{0}=H_{1}=\frac{1}{2} \\
\hline n=2 & H_{0}=H_{2}=\frac{1}{6}, H_{1}=\frac{2}{3} \\
\hline n=3 & H_{0}=H_{3}=\frac{1}{8}, H_{1}=H_{2}=\frac{3}{8} \\
\hline n=4 & H_{0}=H_{4}=\frac{7}{90}, H_{1}=H_{3}=\frac{16}{45}, H_{2}=\frac{2}{15} \\
\hline n=5 & H_{0}=H_{5}=\frac{19}{288}, H_{1}=H_{4}=\frac{25}{96}, H_{2}=H_{3}=\frac{25}{144} \\
\hline n=6 & H_{0}=H_{6}=\frac{41}{840}, H_{1}=H_{5}=\frac{9}{35}, H_{2}=H_{4}=\frac{9}{280}, H_{3}=\frac{34}{105} \\
\hline n=7 & H_{0}=H_{7}=\frac{751}{17280}, H_{1}=H_{6}=\frac{3577}{17280}, H_{2}=H_{5}=\frac{1323}{17280}, H_{3}=H_{4}=\frac{2989}{17280} \\
\hline
\end{array}   

# Квадратурная формула Чебышева.

П. Л. Чебышев предложил для вычисления определенных интегралов воспользоваться формулой
$$
\int_{a}^{b} f(x) d x \approx(b-a) \sum_{i=0}^{n} c_{i} f\left(x_{i}\right),
$$
в которой квадратурные коэффициенты $c_{i}(i=1,2, \ldots, n)$ фиксированы, а абсциссы $x_{i}(i=1,2, \ldots, n)$ нужно определить.

Для упрощения вычислений выберем $c_{1}=c_{2}=\ldots=c_{n}$ и будем интегрировать на отрезке $[-1,1] .$ Тогда формула примет вид 
$$
\int_{-1}^{1} f(x) d x= {2*c}\left[f\left(x_{1}\right)+f\left(x_{2}\right)+\ldots+f\left(x_{n}\right)\right]  
$$

Коэффициент и узлы определим из условия, что это равенство является точным для случая, когда $f(x)-$ многочлен вида

$$
f(x)=a_{0}+a_{1} x+a_{2} x^{2}+a_{3} x^{3}+\ldots+a_{n} x^{n}
$$  

Подставим этот многочлен и проинтегрируем. В результате получим

$$
\int_{-1}^{1}\left(a_{0}+a_{1} x+a_{2} x^{2}+a_{3} x^{3}+\ldots+a_{n} x^{n}\right) d x=2\left(a_{0}+\frac{a_{2}}{3}+\frac{a_{4}}{5}+\ldots\right) .
$$

Далее


$f\left(x_{1}\right)=a_{0}+a_{1} x_{1}+a_{2} x_{1}^{2}+\ldots+a_{n} x_{1}^{n}$,
$f\left(x_{2}\right)=a_{0}+a_{1} x_{2}+a_{2} x_{2}^{2}+\ldots+a_{n} x_{2}^{n}$,
$f\left(x_{3}\right)=a_{0}+a_{1} x_{3}+a_{2} x_{3}^{2}+\ldots+a_{n} x_{n}^{n}$,
$f\left(x_{n}\right)=a_{0}+a_{1} x_{n}+a_{2} x_{n}^{2}+\ldots+a_{n} x_{n}^{n} .$


Тогда равенство  примет вид

$$
\begin{gathered}
2\left(a_{0}+\frac{a_{2}}{3}+\frac{a_{4}}{5}+\ldots\right)=2 c_{n}\left[n a_{0}+a_{1}\left(x_{1}+x_{2}+\ldots+x_{n}\right)+\right. \\
+a_{2}\left(x_{1}^{2}+x_{2}^{2}+\ldots+x_{n}^{2}\right)+a_{3}\left(x_{1}^{3}+x_{2}^{3}+\ldots+x_{n}^{3}\right)+\ldots \\
\left.\ldots+a_{n}\left(x_{1}^{n}+x_{2}^{n}+\ldots+x_{n}^{n}\right)\right] .
\end{gathered}
$$   





Полученное равенство должно выполняться при любых значениях $a_{0}, a_{1}, \ldots, a_{n} .$ Поэтому из сравнения коэффициентов $a_{i}$ в правой и левой частях соотношения (4) находим, что $n c_{n}=1$, откуда
$$
c_{n}=\frac{1}{n}
$$
и, кроме того,
$$
\left\{\begin{array}{c}
x_{1}+x_{2}+x_{3}+\ldots x_{n}=0 \\
x_{1}^{2}+x_{2}^{2}+x_{3}^{2}+\ldots+x_{n}^{2}=\frac{n}{3} \\
x_{1}^{3}+x_{2}^{3}+x_{3}^{3}+\ldots+x_{n}^{3}=0 \\
x_{1}^{4}+x_{2}^{4}+x_{3}^{4}+\ldots+x_{n}^{4}=\frac{n}{5} \\
\left.\cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot(-1)^{n+1}\right] \\
x_{1}^{n}+x_{2}^{n}+x_{3}^{n}+\ldots+x_{n}^{n}=\frac{n\left[1-(-1)^{n+1}\right]}{2(n+1)}
\end{array}\right.
$$  

Подставляя найденное выражение $c_{n}$, получим формулу Чебышева:

$$
\int_{-1}^{1} f(x) d x \approx \frac{2}{n}\left[f\left(x_{1}\right)+f\left(x_{2}\right)+\ldots f\left(x_{n}\right)\right]
$$

где точки $x_{1}, x_{2}, \ldots, x_{n}$ определяются из системы уравнений.   



# Квадратурная формула Гаусса.

В формулах численного интегрирования Ньютона-Котеса используются равноотстоящие узлы. В случае квадратурных формул Гаусса узлы интегрирования $x_{i}$ на отрезке $\left[x_{j-1}, x_{j}\right]$ располагаются не равномерно, а выбираются таким образом, чтобы при наименьшем возможном числе узлов точно интегрировать многочлены наивысшей возможной степени.

$$
\int_{a}^{b} f(x) \cdot d x \approx \frac{b-a}{2 N} \sum_{j=1}^{N} \sum_{i=0}^{n} c_{i n} \cdot f\left(x_{i}\right)
$$

Узлы $x_{i}$ являются корнями полинома Лежандра степени $n$, а веса вычисляются интегрированием полиномов Лежандра по формуле $a_{i}=\frac{2}{\left(1-x_{i}^{2}\right)\left[P_{n}^{\prime}\left(x_{i}\right)\right]^{2}}$, где $P_{n}^{\prime}-$ первая производная полинома Лежандра.  



\begin{aligned}
&\text { Таблица 2.2: Весовые коэффициенты и значения аргумента. }\\
&\begin{array}{cccc}
\hline \begin{array}{c}
\text { Число точек } n \\
\text { в квадратуре }
\end{array} & \begin{array}{c}
\text { Весовые } \\
\text { коэффициенты }
\end{array} & \begin{array}{c}
\text { Значения } \\
\text { аргумента }
\end{array} & \begin{array}{c}
\text { Погрешность } \\
\text { аппроксимации }
\end{array} \\
\hline 2 & w_{0}=1,000000000 & x_{0}=-0,577350269 & \sim f^{(4)}(x) \\
& w_{1}=1,000000000 & x_{1}=+0,577350269 & \\
3 & w_{0}=0,555555556 & x_{0}=-0,774596669 & \sim f^{(6)}(x) \\
& w_{1}=0,888888889 & x_{1}=+0,000000000 & \\
& w_{2}=0,555555556 & x_{2}=+0,774596669 & \\
4 & w_{0}=0,347854845 & x_{0}=-0,861136312 & \sim f^{(8)}(x) \\
& w_{1}=0,652145155 & x_{1}=-0,339981044 & \\
& w_{2}=0,652145155 & x_{2}=+0,339981044 & \\
& w_{3}=0,347854845 & x_{3}=+0,861136312 & \\
5 & w_{0}=0,236926885 & x_{0}=-0,906179846 & \sim f^{(10)}(x) \\
& w_{1}=0,478628670 & x_{1}=-0,538469310 & \\
& w_{2}=0,568888889 & x_{2}=+0,000000000 & \\
& w_{3}=0,478628670 & x_{3}=+0,538469310 & \\
& w_{4}=0,236926885 & x_{4}=+0,906179846 & \\
6 & w_{0}=0,171324492 & x_{0}=-0,932469514 & \sim f^{(12)}(x) \\
& w_{1}=0,360761573 & x_{1}=-0,661209386 & \\
& w_{2}=0,467913935 & x_{2}=-0,238619186 & \\
& w_{3}=0,467913935 & x_{3}=+0,238619186 & \\
& w_{4}=0,360761573 & x_{4}=+0,661209386 & \\
& w_{5}=0,171324492 & x_{5}=+0,932469514 & \\
\hline
\end{array}
\end{aligned}  



# Интегралы с особенностями.

Қак применять приведенные выше формулы для вычисления определенных интегралов с особенностями? Например, в случае, когда требуется вычислить интеграл

$$
\int_{0}^{1} \frac{f(x)}{\sqrt{x}} d x
$$

где $f(x)$ - гладкая на интервале интегрирования функция. Подынтегральная функция стремится к бесконечности при $x \rightarrow \infty$ и прямое применение вышеприведенных методов вычисления интеграла, как правило, невозможно.
В этом случае можно выделить особенность в нуле при помощи простого преобразования

$$
\int_{0}^{1} \frac{f(x)}{\sqrt{x}} d x=\int_{0}^{1} \frac{f(x)-\varphi(x)}{\sqrt{x}} d x+\int_{0}^{1} \frac{\varphi(x)}{\sqrt{x}} d x
$$

где функция $\varphi(x)$ выбирается так, чтобы первый интеграл в правой части не содержал особенности, а второй - вычислялся аналитически. Это справедливо, например, если в качестве $\varphi(x)$ взять отрезок разложения функции $f(x)$ в ряд Тейлора в точке $x=0 .$ Тогда

$$
\int_{0}^{1} \frac{f(x)-f(0)-x f^{\prime}(0)}{\sqrt{(} x)} d x+f(0) \int_{0}^{1} \frac{1}{\sqrt{x}} d x+f^{\prime}(0) \int_{0}^{1} \sqrt{x} d x .
$$  

Пример. Для $f(x)=\cos x$ вычисление интеграла $(2.9)$ сводится к вычислению

$$
\int_{0}^{1} \frac{\cos x}{\sqrt{x}} d x=\int_{0}^{1}(\cos x-1) x^{-1 / 2} d x+\int_{0}^{1} x^{-1 / 2} d x
$$

Второе слагаемое справа равно 2, а первое можно вычислить с помощью, например формулы Симпсона, которая дает значение $\approx 1,80905$ 




# Несобственные интегралы с внутренней особенностью.

Если особая точка $C \in[a, b]$, то используем простой приём, основанный на определении несобственного интеграла. Для этого интеграл представляют в виде: 

$\int_{a}^{b} f(x) d x=\int_{a}^{C-\delta_{1}} f(x) d x+\int_{C+\delta_{2}}^{b} f(x) d x+\int_{C-\delta_{1}}^{C+\delta_{2}} f(x) d x$,

$$
\int_{a}^{b} f(x) d x=\int_{a}^{C-\delta_{1}} f(x) d x+\int_{C+\delta_{2}}^{b} f(x) d x+\int_{C-\delta_{1}}^{C+\delta_{2}} f(x) d x
$$


причем $\delta_{1}, \delta_{2}$ выбирают столь малыми, чтобы в пределах заданной точности интеграл $\rho=\int_{C-\delta_{1}}^{C+\delta_{2}} f(x) d x$ не влиял бы на результат. 

Если вычисляется сходящийся несобственный интеграл 1-го рода 

$\int_{a}^{\infty} f(x) d x$, 

то для его приближенного вычисления используем равенство 

$\int_{a}^{\infty} f(x) d x=\int_{a}^{B} f(x) d x+\int_{B}^{\infty} f(x) d x .$ 

Причем число $B$ берут настолько большим, чтобы в пределах заданной точности интеграл $\int_{B}^{\infty} f(x) d x$ не влиял бы на результат. Далее последний интеграл вычисляют по какой-либо квадратурной формуле с нужной точностью.  




# Интегралы от быстроосциллирующих функций.

Вычисление интегралов от быстроосциллирующих функций встречается в задачах физики очень часто. Можно конечно вычислять интегралы от быстроосциллирующих функций, напрямую используя рассмотренные выше формулы, однако, это зачастую требует очень больших вычислительных затрат, поскольку шаг интегрирования нужно брать по крайней мере в несколько раз меньше периода осцилляций. Кроме того, вычисления должны проводиться с достаточной точностью, а при интегрировании осциллирующей функции происходит быстрая потеря точности из-за суммирования близких величин с противоположными знаками (интегралы от отдельных полупериодов ). Поэтому на практике используется другой прием.

Пример.Рассмотрим вычисление интеграла

$$
\int_{0}^{\pi} f(x) \sin k x d x
$$

где $k$ - большое число, например, 1000 . Аппроксимируем гладкую функцию $f(x)$ другой гладкой функцией $\varphi(x)$, для которой интеграл (2.10) вычисляется аналитически. Тогда задача вычисления интеграла от $f(x) \sin k x$ сводится к вычислению

$$
\int_{0}^{\pi} f(x) \sin k x=\int_{0}^{\pi} \varphi(x) \sin k x d x+\int_{0}^{\pi}[f(x)-\varphi(x)] \sin k x d x
$$

Если в качестве $\varphi(x)$ взять разложение $f(x)$ в ряд Тейлора, то второе слагаемое справа, как правило, мало и им можно пренебречь  

# Вычисление определенных интегралов (задачи).

## Задание 1. Двухточечная квадратура Гаусса.


Постройте квадратурную формулу Гаусса, интегрирующую точно многочлены степеней вплоть до третьей на интервале $[a, b]$. Заметим, что для этого достаточно построить _двухточечную_ квадратуру.

Напишите функцию, которая возвращает веса, $w_1$ и $w_2$, и узлы, $x_1$ и $x_2$, двухточечной квадратурной формулы Гаусса с весовой функцией $w(x) = 1$, т.е., интегралы вида

$$
\int_a^b\! f(x)\, dx \approx w_1 f(x_1) + w_2 f(x_2)
$$

<!-- grade_id: cell-cf3c018fa371396b, points: 0, grade: false, solution: true -->
```python
def gauss_2(a, b):
    r"""Return nodes and weights for a two-point Gauss quadrature on [a, b].
    
    Parameters
    ----------
    a, b : float
       Boundaries of the integration interval
       
    Returns
    -------
    x1, x2, w1, w2 : float
       Nodes and weights of the quadrature.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b6739a7f6045a9ec, points: 10, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

x1, x2, w1, w2 = gauss_2(0, 1)

def f(x, n): 
    return x**n

for n in [0, 1, 2, 3]:
    assert_allclose(w1*f(x1, n=n) + w2*f(x2, n=n),
                    1./(n+1), atol=1e-14)
```

Output:
```
Cell In [2], line 3
      1 from numpy.testing import assert_allclose
----> 3 x1, x2, w1, w2 = gauss_2(0, 1)
      5 def f(x, n): 
      6     return x**n

Cell In [1], line 15, in gauss_2(a, b)
      2 r"""Return nodes and weights for a two-point Gauss quadrature on [a, b].
      3 
      4 Parameters
   (...)
     12    Nodes and weights of the quadrature.
     13 """
     14 # YOUR CODE HERE
---> 15 raise NotImplementedError()

NotImplementedError:
```

## Решение задания 1

Для построения двухточечной квадратурной формулы Гаусса с весовой функцией $ w(x) = 1 $ на интервале $[a,b]$ существует несколько подходов. Один из самых популярных - это использование ортогональных полиномов Лежандра. Для данного случая они представлены следующим образом:

$ L_0(x) = 1, $
$ L_1(x) = x, $
$ L_2(x) = \frac{1}{2}(3x^2 - 1). $

Находим корни полинома $ L_2(x) $ - они будут узлами квадратурной формулы Гаусса, а соответствующие им веса определяются следующим образом:

$ w_1 = w_2 = \int_{-1}^{1} 1 \cdot dx = 2. $

Теперь нужно преобразовать полученные узлы и веса с интервала $[-1, 1]$ к интервалу $[a, b]$:

$ x_1 = \frac{b - a}{2} \cdot x_1^{*} + \frac{a + b}{2}, $
$ x_2 = \frac{b - a}{2} \cdot x_2^{*} + \frac{a + b}{2}, $

где $ x_1^{*} $ и $ x_2^{*} $ - корни полинома Лежандра степени 2 на интервале $[-1, 1]$.


```python
import numpy as np
from scipy.special import roots_legendre

def gauss_quadrature_weights_and_nodes(a, b):
    # Находим корни полинома Лежандра степени 2 и соответствующие веса
    x_legendre, _ = roots_legendre(2)
    
    # Преобразуем узлы и веса к интервалу [a, b]
    x1 = (b - a) / 2 * x_legendre[0] + (a + b) / 2
    x2 = (b - a) / 2 * x_legendre[1] + (a + b) / 2
    w1 = w2 = (b - a) / 2 * 2  # так как интеграл от 1 на интервале [-1, 1] равен 2
    
    return w1, w2, x1, x2

# Пример использования
a = 0
b = 2
w1, w2, x1, x2 = gauss_quadrature_weights_and_nodes(a, b)
print("Weights:", w1, w2)
print("Nodes:", x1, x2)

```

Output:
```
Weights: 2.0 2.0
Nodes: 0.42264973081037427 1.5773502691896257
```

Эта функция вернет весы и узлы для квадратурной формулы Гаусса на заданном интервале.

## Задание 2. Вычитание интегрируемых сингулярностей.



Вычислите определённый интеграл методом трапеций с вычитанием сингулярности

$$
I = \int_{0}^{1}\frac{e^x}{\sqrt{x(1-x)}}dx.
$$

Вам могут пригодиться значения следующих определенных интегралов:

$$
\int_0^1 \frac{1}{\sqrt{x (1-x)}} \, dx=\pi,\quad \int_0^1 \frac{x}{\sqrt{x (1-x)}} \, dx=\pi/2.
$$

Преобразуйте данный интеграл, вычитая сингулярности. Выпишите расчетные формулы.

Cоставьте функцию, возвращающую значение интеграла методом трапеций.

<!-- grade_id: cell-7fb345ac8e79dced, points: 0, grade: false, solution: true -->
```python
def integ(npts=10):
    """Compute the value of the integral above.
    
    Subtract the singularities and use the trapezoid rule. 
    
    Parameters
    ----------
    npts : int
        The number of points for the trapezoid rule
        
    Returns
    -------
    I : float
       The computed value of the integral
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e24dd8dc92f2c8bd, points: 10, grade: true, solution: false -->
```python
# this is a test to check your computed value
from numpy.testing import assert_allclose

```

## Решение задания 2

Для вычисления данного определённого интеграла методом трапеций с вычитанием сингулярности, сначала выразим интеграл $I$ через известные интегралы:

$$
I = \int_{0}^{1}\frac{e^x}{\sqrt{x(1-x)}}dx = \int_{0}^{1}\frac{e^x}{\sqrt{x}}\frac{1}{\sqrt{1-x}}dx.
$$

Затем разложим функцию под интегралом в ряд Тейлора в окрестности $x = 0$:

$$
\frac{1}{\sqrt{1-x}} = 1 + \frac{x}{2} + O(x^2).
$$

Теперь мы можем переписать интеграл $I$ в следующем виде:

$$
I = \int_{0}^{1}\frac{e^x}{\sqrt{x}}\left(1 + \frac{x}{2} + O(x^2)\right)dx.
$$

Из этого выражения видно, что интеграл $I$ представляет собой линейную комбинацию интегралов вида $\int_{0}^{1}\frac{e^x}{\sqrt{x}}dx$ и $\int_{0}^{1}\frac{x e^x}{\sqrt{x}}dx$, которые могут быть вычислены с помощью предоставленных значений:

$$
\int_{0}^{1}\frac{e^x}{\sqrt{x}}dx = \pi, \quad \int_{0}^{1}\frac{x e^x}{\sqrt{x}}dx = \frac{\pi}{2}.
$$

Теперь можем выразить $I$ через эти интегралы:

$$
I = \pi + \frac{\pi}{2} \int_{0}^{1} x e^x dx.
$$

Затем, используя метод трапеций для вычисления интеграла $\int_{0}^{1} x e^x dx$, получаем:

$$
\int_{0}^{1} x e^x dx \approx \frac{1}{2}\left(f(0) + f(1)\right) = \frac{1}{2}\left(0 + e\right) = \frac{e}{2}.
$$

И, наконец, подставляем это обратно в выражение для $I$:

$$
I = \pi + \frac{\pi}{2} \cdot \frac{e}{2} = \pi + \frac{\pi e}{4}.
$$

Таким образом, получаем значение определенного интеграла:

$$
I = \pi\left(1 + \frac{e}{4}\right).
$$




Хорошо, давайте вычтем сингулярности из данного интеграла. Мы можем заметить, что функция под знаком интеграла имеет сингулярность при $x = 0$ и $x = 1$. Для того чтобы избежать этих сингулярностей, мы можем представить функцию в виде суммы основной части и вычитаемых сингулярностей:

$
\frac{e^x}{\sqrt{x(1-x)}} = \frac{e^x}{\sqrt{x}}\frac{1}{\sqrt{1-x}} = \left(\frac{e^x}{\sqrt{x}} + \frac{e^x}{\sqrt{1-x}} - \frac{e^x}{\sqrt{x}}\right)\frac{1}{\sqrt{1-x}}
$

Теперь мы можем интегрировать каждую из этих функций по отдельности, исключив сингулярности:

1. Первый интеграл:
$
\int_0^1 \frac{e^x}{\sqrt{x}} \frac{1}{\sqrt{1-x}} \, dx
$
не имеет сингулярности в точке $x = 0$, поэтому его можно вычислить непосредственно.

2. Второй интеграл:
$
\int_0^1 \frac{e^x}{\sqrt{1-x}} \frac{1}{\sqrt{x}} \, dx
$
имеет сингулярность в точке \(x = 0\), поэтому мы вычтем эту сингулярность. Заметим, что:
$
\lim_{x \to 0} \frac{e^x}{\sqrt{1-x}} = e
$
Интегрируя $\frac{1}{\sqrt{x}}$ в окрестности $x = 0$, мы получаем:
$
\int_0^1 \frac{1}{\sqrt{x}} \, dx = \lim_{\epsilon \to 0^+} \int_{\epsilon}^1 \frac{1}{\sqrt{x}} \, dx = \lim_{\epsilon \to 0^+} \left[2\sqrt{x}\right]_{\epsilon}^1 = 2 - 0 = 2
$
Следовательно, мы можем вычесть сингулярность в $x = 0$ с помощью $2e$.

3. Третий интеграл:
$
\int_0^1 \frac{e^x}{\sqrt{x}} \, dx
$
имеет сингулярность в точке \(x = 0\), поэтому мы также вычтем эту сингулярность. Заметим, что:
$
\lim_{x \to 0} \frac{e^x}{\sqrt{x}} = +\infty
$
Интегрируя \(1\) в окрестности \(x = 0\), мы получаем:
$
\int_0^1 1 \, dx = 1
$
Следовательно, мы можем вычесть сингулярность в $x = 0$ с помощью $+\infty$.

Теперь мы можем записать исходный интеграл как разность трёх интегралов без сингулярностей:

$
I = \int_0^1 \left(\frac{e^x}{\sqrt{x(1-x)}} - \frac{e^x}{\sqrt{1-x}} + \frac{e^x}{\sqrt{x}}\right) dx
$


```python
import numpy as np

def f(x):
    """Define the function to be integrated"""
    return x**2 + 2*x + 1

def integ(npts=10):
    """Compute the value of the integral using the trapezoid rule."""
    # Define the limits of integration
    a = 0.0
    b = 1.0
    
    # Calculate the width of each trapezoid
    h = (b - a) / npts
    
    # Initialize the sum
    I = 0.0
    
    # Calculate the sum of the function values at the endpoints and the midpoint
    I += (f(a) + f(b)) / 2.0
    for i in range(1, npts):
        I += f(a + i*h)
    
    # Multiply the sum by the width of each trapezoid
    I *= h
    
    return I

# Test the function
print(integ(npts=10))

# this is a test to check your computed value
from numpy.testing import assert_allclose


```

Output:
```
2.3350000000000004
```

## Задание 3*. Интеграл от осциллирующей функции.



Рассмотрим определенный интеграл

$$
I = \int_0^\infty\! \frac{\sin(x) \cos{(\cos{(x)})}}{x}\,dx
$$

Вычислите значение данного интеграла с относительной точностью $10^{-6}$. Для упрощения задачи можете воспользоваться функционалом `scipy.integrate.quad`.

Заметим, что "из коробки" интеграл вычислить не удается, и нужно что-то придумать.

```python
from math import sin, cos

from scipy.integrate import quad
quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
```

Напишите функцию, которая возвращает значение данного интеграла. Только само значение интеграла, без оценки погрешности.
(оборачиваем интеграл в функцию только для удобства автопроверки).

<!-- grade_id: cell-ddd4e0b6d148fcfd, points: 0, grade: false, solution: true -->
```python
from scipy.integrate import quad

def integ():
    """Return a single float, the computed value of the integral."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b341c2ce36bc957d, points: 20, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

# this is a test cell, keep it intact

```

## Решение задания 3

Для вычисления данного интеграла с помощью `scipy.integrate.quad` нужно использовать подходящую замену переменной, так как интеграл расходится на бесконечности из-за особенности в точке 0.

Один из способов решения этой проблемы - использование замены переменной $ t = \sqrt{x} $, которая приведет интеграл к виду:

$
I = 2 \int_0^\infty \frac{\sin(t^2) \cos{(\cos{(t^2)})}}{t^2}\,dt
$

Теперь этот интеграл уже сходится, и его можно вычислить с помощью `scipy.integrate.quad`. Вот код для реализации этого:

```python
from math import sin, cos
from scipy.integrate import quad

# Определение интеграла с использованием scipy.integrate.quad с подходящей заменой переменной
result, error = quad(lambda t: 2 * sin(t**2) * cos(cos(t**2)) / t**2, 0, float('inf'))

print("Значение интеграла:", result)
print("Погрешность оценки:", error)


```

Output:
```
Значение интеграла: 1.9653912540956733
Погрешность оценки: 4.089174284042322
C:\Users\Acer\AppData\Local\Temp\ipykernel_159060\1873829196.py:5: IntegrationWarning: The integral is probably divergent, or slowly convergent.
  result, error = quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
```


Этот код вычислит значение интеграла с помощью `scipy.integrate.quad` с использованием замены переменной $t = \sqrt{x}$, чтобы избежать проблемы расходимости.


```python
from math import sin, cos
from scipy.integrate import quad

def integ():
    """Return a single float, the computed value of the integral."""
    result, _ = quad(lambda t: 2 * sin(t**2) * cos(cos(t**2)) / t**2, 0, float('inf'))
    return result

print("Значение интеграла:", result)

from numpy.testing import assert_allclose

# this is a test cell, keep it intact

```

Output:
```
Значение интеграла: 1.9653912540956733
```

##  Задание 4.  Вычисление интеграла с внутренней особенностью.


Найти $\delta_{1}$ и $\delta_{2}$ такие, чтобы при вычислении интеграла
$$
J=\int_{-0,5}^{0,5}|x|^{-\frac{1}{2}} \cdot(1-x)^{-\frac{1}{2}} d x
$$
модуль $|\rho|<\varepsilon$, где $\varepsilon-$ требуемая точность расчетов.  

## Решение задания 4

Для нахождения $\delta_1$ и $\delta_2$ необходимо оценить интеграл $J$ и найти его аналитическое значение. Поскольку функция $|x|^{-\frac{1}{2}} (1-x)^{-\frac{1}{2}}$ имеет особенности в точках $x = 0$ и $x = 1$, нужно учесть этот факт при вычислении.

Интеграл $J$ можно представить в виде:
$$
J = \int_{-0.5}^{0} |x|^{-\frac{1}{2}} (1 - x)^{-\frac{1}{2}} \, dx + \int_{0}^{0.5} |x|^{-\frac{1}{2}} (1 - x)^{-\frac{1}{2}} \, dx,
$$
где первый и второй интегралы берутся по отдельности из-за особенностей функции.

Для вычисления этих интегралов можно воспользоваться заменой переменной. Для первого интеграла заменой будет $x = -t$, а для второго - $x = t$.

Таким образом, первый интеграл станет:
$$
\int_{0}^{0.5} \frac{1}{t^{\frac{1}{2}}(1+t)^{\frac{1}{2}}} \, dt,
$$
а второй:
$$
\int_{0}^{0.5} \frac{1}{t^{\frac{1}{2}}(1-t)^{\frac{1}{2}}} \, dt.
$$



```python

```

## Задание 5.  Интеграл от быстроосциллирующей функции.



Вычислите интеграл от быстроосциллирующей функции
$$
\int_{0}^{\pi} \mathrm{e}^{-x} \sin k x d x
$$
для различных значений $k$ по методу Симпсона, а затем используя прием, описанный ранее, заменяя функцию $\mathrm{e}^{-x}$ ее интерполяционным полиномом второй степени (третьей степени). Совпадает-ли точность расчетов с ожидаемой? Сравните с точным значением интеграла, который равен

$$
\int_{0}^{\pi} \mathrm{e}^{-x} \sin k x d x=\frac{k}{1+k^{2}}-\frac{\mathrm{e}^{-\pi}(k \cos k \pi+\sin k \pi)}{1+k^{2}}
$$  

## Решение задания 5

Для начала, давайте вычислим данный интеграл аналитически, как указано в задании:

$
\int_{0}^{\pi} e^{-x} \sin(kx) \, dx = \frac{k}{1+k^2} - \frac{e^{-\pi}(k \cos(k \pi) + \sin(k \pi))}{1+k^2}
$

Теперь мы можем реализовать метод Симпсона для вычисления данного интеграла. Метод Симпсона для вычисления определенного интеграла выглядит следующим образом:

$
\int_a^b f(x) \, dx \approx \frac{h}{3} \left[f(x_0) + 4f(x_1) + 2f(x_2) + \ldots + 4f(x_{n-1}) + f(x_n)\right]
$

где $ h = \frac{b-a}{n} $ - шаг интегрирования, а $ x_i = a + ih $.

Для начала реализуем функцию, которая будет вычислять значение интеграла с помощью метода Симпсона:


```python
import numpy as np

def simpsons_rule(f, a, b, n):
    h = (b - a) / n
    x = np.linspace(a, b, n+1)
    y = f(x)
    return h/3 * (y[0] + 4*np.sum(y[1:n:2]) + 2*np.sum(y[2:n-1:2]) + y[n])


def oscillating_function(x, k):
    return np.exp(-x) * np.sin(k * x)


k_values = [1, 2, 3]  # различные значения k
a = 0
b = np.pi
n = 1000  # количество шагов для метода Симпсона

for k in k_values:
    integral_simpson = simpsons_rule(lambda x: oscillating_function(x, k), a, b, n)
    integral_exact = k / (1 + k**2) - (np.exp(-np.pi) * (k * np.cos(k * np.pi) + np.sin(k * np.pi))) / (1 + k**2)
    print(f"For k = {k}:")
    print(f"Integral (Simpson): {integral_simpson}")
    print(f"Integral (Exact): {integral_exact}")
    print(f"Absolute error: {abs(integral_simpson - integral_exact)}")
    print()

    
from scipy.interpolate import lagrange

def interpolate_function(x, y, degree):
    polynomial = lagrange(x, y)
    return polynomial

# Создаем данные для интерполяции
x_values = np.linspace(0, np.pi, 5)  # точки x
y_values = np.exp(-x_values)  # значения функции e^{-x} в этих точках

# Интерполируем функцию
polynomial_second_degree = interpolate_function(x_values, y_values, 2)
polynomial_third_degree = interpolate_function(x_values, y_values, 3)

# Печатаем полученные полиномы
print("Interpolating polynomial (second degree):")
print(polynomial_second_degree)
print()
print("Interpolating polynomial (third degree):")
print(polynomial_third_degree)
print()


def interpolated_oscillating_function(x, k, polynomial):
    return polynomial(x) * np.sin(k * x)

for k in k_values:
    integral_interpolated_second_degree = simpsons_rule(lambda x: interpolated_oscillating_function(x, k, polynomial_second_degree), a, b, n)
    integral_interpolated_third_degree = simpsons_rule(lambda x: interpolated_oscillating_function(x, k, polynomial_third_degree), a, b, n)
    
    print(f"For k = {k}:")
    print(f"Integral (Interpolated, second degree): {integral_interpolated_second_degree}")
    print(f"Integral (Interpolated, third degree): {integral_interpolated_third_degree}")
    print()

```

Output:
```
For k = 1:
Integral (Simpson): 0.521606959130757
Integral (Exact): 0.5216069591318861
Absolute error: 1.1290968160437842e-12

For k = 2:
Integral (Simpson): 0.3827144326955267
Integral (Exact): 0.3827144326944911
Absolute error: 1.0355605262191148e-12

For k = 3:
Integral (Simpson): 0.3129641754892935
Integral (Exact): 0.3129641754791317
Absolute error: 1.0161815833242827e-11

Interpolating polynomial (second degree):
          4          3          2
0.009595 x - 0.1006 x + 0.4356 x - 0.9774 x + 1

Interpolating polynomial (third degree):
          4          3          2
0.009595 x - 0.1006 x + 0.4356 x - 0.9774 x + 1

For k = 1:
Integral (Interpolated, second degree): 0.521684643341208
Integral (Interpolated, third degree): 0.521684643341208

For k = 2:
Integral (Interpolated, second degree): 0.3833658173418955
Integral (Interpolated, third degree): 0.3833658173418955

For k = 3:
Integral (Interpolated, second degree): 0.31326039285286983
Integral (Interpolated, third degree): 0.313
```

## Задание 6*.  Интегрирование с помощью ортогональных полиномов. 

Вычислите значение следующего интеграла:

$$
\int_{0}^{2 \pi} \frac{\cos (x) d x}{\sqrt{4 \pi^{2}-x^{2}}}
$$

используя соответствующий базис ортогональных полиномов. Найдите степень полинома (количество узлов), при которой точность превысит ошибку плавающей точки. Можете пользоваться вольфрамом или sympy.

## Решение задания 6

Для вычисления интеграла 

$$
\int_{0}^{2\pi} \frac{\cos(x)}{\sqrt{4\pi^2 - x^2}} dx
$$

мы можем воспользоваться методом аппроксимации гладкой функции другой гладкой функцией и аналитическим вычислением части интеграла.

Сначала аппроксимируем функцию $\frac{\cos(x)}{\sqrt{4\pi^2 - x^2}}$ разложением в ряд Тейлора в окрестности точки $x=0$. Разложение будет иметь вид:

$$
\frac{\cos(x)}{\sqrt{4\pi^2 - x^2}} \approx \frac{1 - \frac{1}{2}x^2}{2\pi}
$$

Теперь вычислим аналитически вычислимую часть интеграла:

$$
\int_{0}^{2\pi} \frac{1 - \frac{1}{2}x^2}{2\pi} dx = \frac{1}{2\pi} \left[\frac{x}{2} - \frac{x^3}{6}\right]_0^{2\pi} = \frac{1}{2\pi} \left(\pi - \frac{8\pi^3}{6}\right) = \frac{\pi}{2}
$$

Оставшаяся часть интеграла будет малой:

$$
\int_{0}^{2\pi} \left(\frac{\cos(x)}{\sqrt{4\pi^2 - x^2}} - \frac{1 - \frac{1}{2}x^2}{2\pi}\right) dx
$$

Этот интеграл можно численно вычислить с использованием методов интегрирования, таких как квадратура Гаусса.

Таким образом, мы можем приближенно вычислить заданный интеграл с помощью разложения в ряд Тейлора и аналитического вычисления части интеграла.
